---
title: 'Local Outlier Factor (LOF) for anomaly detection'
description: 'In this chapter, we will explore the use of Local Outlier Factor (LOF) for detection, calculate LOF scores, and visualize them. '
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Visualizing LOF Scores

```yaml
type: VideoExercise
key: 92f6ccf829
xp: 50
```

`@projector_key`
29d3c8842a2de2c8d14b9bbc23104ad5
